## Exercices avant le TP4

## Génération de listes aléatoires

Pour comparer les algorithmes de tri, nous aurons besoin de générer des listes aléatoires. Pour tirer un entier au hasard, nous pourrons utiliser la fonction `randint` de `random`, comme ci-dessous, qui renvoie un entier compris (au sens large) entre les deux bornes données.

```python
from random import randint
randint(1,10)
```

Output:
```
5
```

Écrire une fonction `randListe(n, a, b)` qui renvoie une liste de longueur `n` dont les valeurs ont été tirées au hasard dans l'intervalle d'entiers de `a` à `b`. Par exemple, `randListe(10, 1, 20)` pourrait retourner une liste comme `[7, 9, 10, 15, 20, 13, 11, 7, 13, 11]`.

<!-- grade_id: cell-4e5e2cfdc152ca55, points: 0, grade: false, solution: true -->
```python
def randListe(n, a, b):
    ### BEGIN SOLUTION
    r = []
    for i in range(n):
        r.append(randint(a, b))
    return r
    ### END SOLUTION
```

```python
randListe(10, 1, 20)
```

Output:
```
[15, 4, 16, 3, 10, 5, 7, 3, 3, 1]
```

<!-- grade_id: cell-38dc4a6cc23a3d70, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
for (n,a,b) in  [(13,7,19), (1,3,3), (2,0,1), (12,6,123)]:
    lnab = randListe(n,a,b)
    assert len(lnab) == n
    assert min(lnab) >= a
    assert max(lnab) <= b
### END HIDDEN TESTS
```

## Mesure du temps d'exécution

Pour mesurer le temps d'exécution, nous allons utiliser `perf_counter`. Ci-dessous, nous l'utilisons pour mesurer le temps nécessaire pour la création d'une liste aléatoire de longueur `100000`.

```python
from time import perf_counter
```

```python
t0 = perf_counter()
L = randListe(100000, 1, 10)
t1 = perf_counter()
t1 - t0
```

Output:
```
0.08004040003288537
```

Écrire une fonction `listeTemps(n)` qui va boucler `n` fois, à chaque fois créant une liste de nombres de longueur `100000`, mesurant le temps nécessaire pour *chaque* création, puis renverra la liste des `n` temps ainsi obtenus.

```python
def listeTemps(n):
    temps = []
    for _ in range(n):
        t0 = perf_counter()
        L = randListe(100000, 1, 10)
        t1 = perf_counter()
        temps.append(t1 - t0)
    return temps
```

```python
listeTemps(20)
```

Output:
```
[0.0818862000014633,
 0.07699830003548414,
 0.06481950008310378,
 0.056164100067690015,
 0.05199569999240339,
 0.05159029993228614,
 0.05261629994492978,
 0.052938199951313436,
 0.05538270005490631,
 0.05277329997625202,
 0.05084299994632602,
 0.05123199999798089,
 0.05133299995213747,
 0.05260019993875176,
 0.06269019993487746,
 0.05225760000757873,
 0.051493900013156235,
 0.05111699993722141,
 0.051631500013172626,
 0.051270099938847125]
```

## Fusion de deux listes triées

Implémenter le pseudo-code vu en cours pour la fusion de deux listes triées: définir une fonction `fusion(A, B)`. Par exemple, si `A = [1, 3, 4, 8, 12, 14, 19]` et `B = [2, 6, 9, 10, 11, 15, 20, 40]`, votre fonction doit renvoyer la liste `[1, 2, 3, 4, 6, 8, 9, 10, 11, 12, 14, 15, 19, 20, 40]`.

## Tri à bulles

Implémenter le pseudo-code du cours pour le tri à bulles.